<div class="alert alert-info">
Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`, as well as your name and student id below.

Make sure you execute **all** executable cells. Remember that you can use `SHIFT+ENTER` to execute a cell. It will install a testing framework used by the notebook and executing this is necessary to make sure that the notebook functions properly.

Before submitting do not forget to **save** your work (`CTRL+S`, or `File/Save` in the file menu).
</div>

```python
NAME = "Kristopher Quaife-Glasier"
STUDENT_ID = "1516240"
```

If you are running the notebook on your own computer, you should install nbgrader and nose manually from the command line. The next line assumes that you are using the VM we provided: 'cmput274' is the password passed to `sudo`.

```python
# %%bash
# echo 'cmput274' | sudo -HS pip3 install nbgrader
# echo 'cmput274' | sudo -HS pip3 install nose
```

---

# Kruskal's vs. Prim's algorithm

## Part 1
Find a minimum spanning tree (MST) in the graph shown here by hand using Kruskal's algorithm.
<img src="graph.svg">
<div class="alert alert-success">
In the next cell, assign to the variable `mst` a **list** that contains the edges of a MST of this graph in the order that they are added. Use the `edge` function to specify edges.</div>

<!-- grade_id: msthand_answer, points: 0, grade: false, solution: true -->
```python
def edge(x,y):
    return frozenset({x,y})
mst = [edge(1,4), edge(1,0), edge(0,2), edge(2,5), edge(5,6), edge(6,8), edge(8,9), edge(9,3), edge(3,7)]

```

<!-- grade_id: msthand_test, points: 3, grade: true, solution: false -->
```python
%reload_ext autoreload
%aimport test
%aimport kruskal
from nose.tools import assert_equal

# not cyclic
assert not kruskal.is_cyclic(mst)
# covers all vertices
assert frozenset.union(*mst)==frozenset(range(10))

# check total weight:
weights = \
    {edge(0,1):2,edge(0,2):3,edge(1,2):3,edge(1,3):3,edge(1,4):1,edge(2,5):2,
     edge(3,4):4,edge(4,5):4,edge(3,7):2,edge(3,9):1,edge(4,6):3,edge(5,6):1,
     edge(5,8):5,edge(6,9):4,edge(6,8):3,edge(7,9):2,edge(8,9):2}    
total = sum([weights[e] for e in mst])
assert total==17

print("Passed public test")
```

Output:
```
Passed public test
```

<div class="alert alert-info">
Your answer is correct if no error is raised when you execute the next cell and if the cell prints `Passed public test`. On an exam some tests may be hidden from you, which will also be taken into account whem marking your solution.</div>

## Part 2
The following is Prim's algorithm for computing a minimum spanning tree of a connected graph $G = (V,E)$ with edge weights $w_e \geq 0, e \in E$.


1. Let $v$ be any arbitrarily chosen vertex.
2. Let $S := \{v\}$ and $T := \{ \}$.
3. **while** $S \neq V$
    1. Among the edges with exactly one endpoint $S$, pick one with minimum weight (break ties arbitrarily). Call this edge $(u,w)$ where $w$ is the endpoint not in $S$.
    2. Update $S := S \cup \{w\}$ and $T = T \cup \{(u,w)\}$.
4. **return** $T$

Address these problems.

### Part 2.1
Trace the execution of this algorithm on the same graph as in problem 1 starting with the leftmost vertex. For convenience the graph is repeated:
<img src="graph.svg">
<div class="alert alert-success">
In the next cell, assign to the variable `prim_trace` a set that contains the edges of a MST of this graph. Use the `edge` function to specify edges.</div>

<!-- grade_id: prim_answer, points: 0, grade: false, solution: true -->
```python
# This is how the solution starts..
prim_trace = [edge(0,1),edge(1,4),edge(4,6),edge(6,5),edge(5,2),edge(6,8),edge(8,9),edge(9,3),edge(3,7)]

```

<!-- grade_id: prim_test, points: 2, grade: true, solution: false -->
```python
# not cyclic
assert not kruskal.is_cyclic(prim_trace)
# covers all vertices
assert frozenset.union(*prim_trace)==frozenset(range(10))
# check total weight
total = sum([weights[e] for e in prim_trace])
assert total==17
print("Passed public test")
```

Output:
```
Passed public test
```

## Food for thought
1. Think of how to implement Prim's algorithm as efficiently as possible. $O(|E| \log |E|)$ is possible using heaps.
2. At the heart of the proof of the correctness of Kruskal's algorithm was the fact that a spanning tree $T$ is a minimum spanning tree if and only if for every edge $e$ not in $T$
and for every edge $f$ lying on the unique cycle in $T \cup \{e\}$, $w(e) \geq w(f)$ holds.
How can you use this fact to demonstrate that Prim's algorithm finds a minimum spanning tree? If it helps make your argument simpler,
you can assume that no two edges have the same weight (though Prim's algorithm does work in general).

# Finishing up

<br>

<div class="alert alert-danger">

When you are done, save the notebook (`CTRL+S`) and the run the whole notebook from the beginning. You can do this by selecting `Kernel/Restart & Run All` from the menu on the top of the screen. This should print `"All tests passed"` for full marks.<br>

Finally, when you are done with all the problems, you can run <br>

`nbgrader validate *.ipynb`<br> 

from the command line while in the directory where this notebook is. This runs all the tests in the notebook again and it is good to run this to double check whether you actually saved your solutions and that you are in the correct directory.<br>
If everything is as expected, run <br>

`./package.sh` <br>

from the command line in the same directory. This creates the file `practice-final.tar.gz`, which you will need to submit through eclass as usual. 
</div>

```python
print("All tests passed")
```

Output:
```
All tests passed
```

---

Scratch space: Add as many cells as you wish here to experiment. You can also add extra cells anywhere in the notebook, even when you are solving some problem.

```python

```